# Problem 1: Linear data regression - Detecting changing climate (*8.5 points*)

For this problem we will again use the NOAA climate data from Exercises 6 and 7 of the [Geo-Python course](https://geo-python.github.io) with a slight twist.
As a reminder, you will be working with a climate data file from the [US National Oceanographic and Atmospheric Administration (NOAA) climate database](https://www.ncdc.noaa.gov/cdo-web/).
In order to jump right into things, we have provided you with [a data file that has already been processed](data/seasonal-temps.csv) containing the means and standard deviations of seasonal temperatures (in Celsius) for winter and summer for the years 1953-2019.
The first 5 lines of the data file are:

```
Year,Winter mean,Winter std,Summer mean,Summer std
1953,-5.600255,6.126864,16.406055,3.934636
1954,-6.244444,6.804128,15.292479,3.890385
1955,-4.430118,6.128557,16.218130,5.347337
1956,-10.640902,8.521185,14.641628,4.101369
```

Your task is to load this data and use a least squares regression to fit trendlines to the temperature data.

*If you're looking for a challenge, you are welcome to load in the [original data file from Exerise 6 (029740.txt)](https://geo-python.github.io/site/notebooks/L6/advanced-data-processing-with-pandas.html) and process it as was done for Exercises 6 and 7 as a starting point for this problem. As a reminder, the original data file has daily mean, minimum, and maximum temperatures for the Helsinki-Vantaa airport weather station from January 1952-October 2019. Be aware that if you process the data yourself you may get slightly different values for the calculations below.*

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable names mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.

**Your score on this problem will be based on following criteria**:

- Reading the in the data file for this exercise
- Properly converting the linear regressions formulas to a Python function
- Calculating the least-squares regression lines for summer and winter, for two different time periods
- Plotting your results
- Including comments that explain what most lines in the code do
- Uploading your notebook to your GitHub repository for this week's exercise

## Part 1: Reading the data file (*1 point*)

Your first task is to read in [the data file for this problem](data/seasonal-temps.csv) and split the data into column arrays.

- For this, you should use the Pandas `read_csv()` funtion to read in the data file [data/seasonal-temps.csv](data/seasonal-temps.csv) into the variable `data`
    - You do not need to skip any rows or specify an index column

<!-- grade_id: problem_1_read_data, points: 0, grade: false, solution: true -->
```python
# Import Pandas and NumPy
import pandas as pd

data = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_read_data_tests, points: 1, grade: true, solution: false -->
```python
# This test should print the first row of the data file
print("First row of the DataFrame:\n", data.head(1))

```

## Part 2: A linear regression function (*1.5 points*)

Your next task in this problem is to create a function that calculates the slopes `B` and *y*-intercepts `A` for [unweighted linear regression lines](https://introqg.github.io/site/lessons/L2/least-squares.html). As you may recall from the lecture slides, `A` and `B` can be calculated as

\begin{equation}
  \Large
  A = \frac{\sum{x^2} \sum{y} - \sum{x} \sum{x y}}{\Delta}
\end{equation}

and

\begin{equation}
  \Large
  B = \frac{N \sum{x y} - \sum{x} \sum{y}}{\Delta},
\end{equation}

where

\begin{equation}
  \Large
  \Delta = N \sum{x^2} - \left(\sum{x} \right)^2.
\end{equation}

- Create a function `linregress` that returns the slope `B` and *y*-intercept `A` for a regression line to data `x` and `y`.
    - **Note**: You can return both `A` and `B` in the same `return` statement by separating the values with a comma

<!-- grade_id: problem_1_function, points: 0, grade: false, solution: true -->
```python
# Import NumPy in case we want to use it
import numpy as np

def linregress(x, y):
    """Returns A, B, the y-intercept and slope of a least squares regression line."""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_function_tests, points: 1.5, grade: true, solution: false -->
```python
# These visible tests with fake data should work
from nose.tools import ok_, assert_equal

# Fake data for linregress tests
x1 = np.array([1.0, 2.0])
y1 = np.array([1.0, 2.0])
x2 = np.array([1.0, 2.0, 2.5, 3.6, 4.2, 5.9])
y2 = np.array([1.0, 2.0, 2.7, 3.1, 4.7, 5.7])

# Fake line slopes, intercepts
a1, b1 = linregress(x1, y1)
a2, b2 = linregress(x2, y2)

# Print results
print("Slope and intercept for dataset 1: {0:.4f} and {1:.4f}.".format(b1, a1))
print("Slope and intercept for dataset 2: {0:.4f} and {1:.4f}.".format(b2, a2))

# Check that the linregress function works
assert_equal(a1, 0.0)
assert_equal(b1, 1.0)
assert_equal(round(a2, 4), 0.0799)
assert_equal(round(b2, 4), 0.975)
```

## Part 3: Calculating seasonal regression lines (*1.5 points*)

Now we can calculate the slopes and y-intercepts for our weather data to determine how summer and winter temperatures have varied over the observation period. We'll do this for two time ranges: The **entire observation period** and the **last 20 years (2000-2019)**.

- Use your `linregress` function to calculate the slopes and intercepts for summer and winter for the time periods mentioned above (the entire range of years, and 2000-2019).
You should name these values as follows:
    - Entire range of years: `A_winter` and `B_winter` for winter, `A_summer` and `B_summer` for summer
    - Last 20 years: `A_winter20` and `B_winter20` for winter, `A_summer20` and `B_summer20` for summer

**Hint**: In this problem, `x` will be your time data, and `y` will be your temperatures.

<!-- grade_id: problem_1_season_calc, points: 0, grade: false, solution: true -->
```python
# Your calculations go below.
# Note: Replace both of the 'None' values with a single call of linregress()
A_winter, B_winter = None, None
A_summer, B_summer = None, None
A_winter20, B_winter20 = None, None
A_summer20, B_summer20 = None, None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_1_season_calc_variable_tests, points: 0.5, grade: true, solution: false -->
```python
# This test should work
print("Winter (all years) slope {0:.4f} and intercept {1:.1f} should be 0.0451 and -94.5.".format(B_winter, A_winter))

```

<!-- grade_id: problem_1_season_calc_value_tests, points: 1, grade: true, solution: false -->
```python
# This test should work
print("Summer (last 20 years) slope {0:.4f} and intercept {1:.1f} should be 0.0028 and 11.1.".format(B_summer20, A_summer20))

```

## Part 4: Plotting your results (*4.5 points*)

Your last coding task for this problem is to create some plots of your seasonal temperature data and their trendlines.
An example of the types of plots you should create are below.
**Note**: You should create separate plots, not use subplots as shown in the example below.

![Temperature trends in Helsinki](img/temp-slopes.png)

- Create 2 plots of the seasonal temperature data: (**1**) one for the entire time covered in the data and (**2**) another for the last 20 years of data
    - Plot the seasonal average temperature data as a Pandas scatter plot with error bars showing the standard deviation in seasonal temperature
    - Also plot the regression lines for each season as well to compare to the yearly temperature data. You may want to consult [the hints for this week's exercise](https://introqg.github.io/qg/lessons/L2/exercise-2.html) about how to plot the regression lines.
    - Plot the temperature data for both seasons on the same plot, using a warm color (e.g., red) for summer and a cool color (e.g., blue) for winter. You can find some convenient color names you can use in the [Matplotlib documentation online](https://matplotlib.org/3.1.0/gallery/color/named_colors.html).
    - Be sure to set the figure size to be slightly larger, as was done in Exercise 1
- Include axis labels, a title and a legend for each plot, as well as text displaying the slope of the regression lines for both the summer and winter seasons
    - If your plot legend is empty, be sure that you have a label included for the scatter plot for each season with the text name of that season
- Add figure captions in the Markdown cell below each Python cell for your plots that describe the plots as if they were in a scientific journal article

<!-- grade_id: plot_all, points: 1.5, grade: true, solution: true -->
```python
# Plot for all years goes here

%matplotlib inline

ax = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: plot_all_caption, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: plot_20, points: 2, grade: true, solution: true -->
```python
# Plot for last 20 years goes here

ax2 = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: plot_20_caption, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

# Problem 2 - Linear correlation: Is a line a good fit for our data? (*6 points*)

In Problem 1 we used least squares regressions to calculate lines that best fit our data. One thing we did not consider, however, was whether a line was a suitable model for our data. To consider this another way, we could ask ourselves how well do our data plot on a line? This question will be the focus of Problem 2, where we will calculate the linear correlation coefficient $r$ to explore how well our data fall on a line. Notably, in doing this we will *not* consider the standard deviations in our seasonal temperatures, only the mean values for the different seasons.

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable names mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.

**Your score on this problem will be based on following criteria**:

- Properly converting the linear correlation coefficient formula to a Python function
- Calculating linear corelation coefficients for summer and winter and for the two different time periods
- Plotting your results
- Including comments that explain what most lines in the code do
- Uploading your notebook to your GitHub repository for this week's exercise

## Part 1: A correlation coefficient function (*1.5 points*)

As we have seen in this week's lesson, the correlation coefficient $r$ can be used to determine how well two variables are linearly correlated with one another. In other words, how well does a line relate the two. Mathematically, we can define the correlation coefficient $r$ as

\begin{equation}
  \Large
  r = \frac{\sum{\left(x_{i} - \bar{x} \right)\left(y_{i} - \bar{y} \right)}}{\sqrt{\sum{\left(x_{i} - \bar{x} \right)^2} \sum{\left(y_{i} - \bar{y} \right)^2}}},
\end{equation}

where $x_{i}$ is the $i$th value along the $x$-axis, $\bar{x}$ is the mean of the values on the $x$-axis, and similarly for the values of $y$.

Your task in this part of the problem is to make a function for calculating $r$.

- Create a function `pearson()` that returns the correlation coefficient $r$ between two variables `x` and `y`

<!-- grade_id: problem_2_function, points: 0, grade: false, solution: true -->
```python
def pearson(x, y):
    """Returns the linear correlation coefficient for two arrays."""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_function_tests, points: 1.5, grade: true, solution: false -->
```python
# These visible tests with fake data should work
from nose.tools import ok_, assert_equal

# Fake data for pearson tests
x1 = np.array([1.0, 2.0, 3.0])
y1 = np.array([1.0, 2.0, 2.1])
x2 = np.array([1.0, 2.0, 2.5, 3.6, 4.2, 5.9])
y2 = np.array([5.7, 4.2, 3.6, 5.4, 2.0, 1.0])

# Fake correlation coefficients
r1 = pearson(x1, y1)
r2 = pearson(x2, y2)

# Print results
print("Correlation coefficient for dataset 1: {0:.4f}.".format(r1))
print("Correlation coefficient for dataset 2: {0:.4f}.".format(r2))

# Check that the pearson function works
assert_equal(round(r1, 4), 0.9042)
assert_equal(round(r2, 4), -0.8156)
```

## Part 2: Calculating correlation coefficients (*1.5 points*)

Next, we can calculate the correlation coefficients for our weather data.

- Calculate correlation coefficients for the summer and winter and for both time ranges (the entire range of years, and 2000-2019) using your `pearson()` function.
You should name these values as follows:
    - Entire range of years: `r_winter` for winter and `r_summer` for summer
    - Last 20 years: `r_winter20` for winter, `r_summer20` for summer

<!-- grade_id: problem_2_r, points: 0, grade: false, solution: true -->
```python
r_winter = None
r_summer = None
r_winter20 = None
r_summer20 = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_2_r_variable_tests, points: 0.5, grade: true, solution: false -->
```python
# This test should work
print('Your winter (all years) correlation coefficient is {0:.4f} and should be 0.3145.'.format(r_winter))

```

<!-- grade_id: problem_2_r_value_test, points: 1, grade: true, solution: false -->
```python
# This test should work
print('Your summer (last 20 years) correlation coefficient is {0:.4f} and should be 0.0148.'.format(r_summer20))

```

## Part 3: Plotting your results, round 2 (*3 points*)

Now we can generate a pair of plots showing the data along with the calculated $r$ values.

- Create 2 plots of the seasonal temperature data: (**1**) one for the entire time covered in the data and (**2**) another for the last 20 years of data
    - Plot the seasonal average temperature data as a Pandas scatter plot **without** error bars for the standard deviation in seasonal temperature
    - Also include the regression lines for each season from Problem 1 to compare to the yearly temperature data. You may want to consult [the hints for this week's exercise](https://introqg.github.io/qg/lessons/L2/exercise-2.html) about how to plot the regression lines.
    - Plot the temperature data for both seasons on the same plot, using a warm color (e.g., red) for summer and a cool color (e.g., blue) for winter. You can find some convenient color names you can use in the [Matplotlib documentation online](https://matplotlib.org/3.1.0/gallery/color/named_colors.html).
    - Be sure to set the figure size to be slightly larger, as was done in Exercise 1
- Include axis labels, a title and a legend for each plot, as well as text displaying the correlation coefficients for both the summer and winter seasons
    - If your plot legend is empty, be sure that you have a label included for the scatter plot for each season with the text name of that season
- Add figure captions in the Markdown cell below each Python cell for your plots that describe the plots as if they were in a scientific journal article

<!-- grade_id: cell-75f0a82446eca8e0, points: 1, grade: true, solution: true -->
```python
# Plot for all years goes here

ax3 = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9bb5a36697760600, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-189ce63a6ebf24c8, points: 1, grade: true, solution: true -->
```python
# Plot for the last 20 years goes here

ax4 = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ee93edb197c1ecc6, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

## Problem 3: Calculating a goodness-of-fit (*3.5 points*)

The final part of the exercise this week is to calculate a goodness-of-fit for the lines fit to the data in Problem 1. The idea here is to consider that we have an expected value of the temperature based on our least-squares regressions, and we can see how well those expected values fit within the standard deviation of the seasonal temperatures.

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable names mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.

**Your score on this problem will be based on following criteria**:

- Properly converting the goodness-of-fit formula to a Python function
- Calculating goodness-of-fit values for summer and winter and for the two different time periods
- Including comments that explain what most lines in the code do
- Uploading your notebook to your GitHub repository for this week's exercise

## Part 1: A goodness-of-fit function (*1.5 points*)

We can start by creating a function to calculate the goodness-of-fit. For this, you can use the reduced chi-squared equation,

\begin{equation}
  \Large
  \chi^{2} = \frac{1}{N} \sum \frac{(O_{i} - E_{i})^{2}}{\sigma_{i}^2}
\end{equation}

where $N$ is the number of ages, $O_{i}$ is the $i$th observed value, $E_{i}$ is the $i$th expected value, and $\sigma_{i}$ is the $i$th standard deviation.

For this part you should:

- Create a function called `chi_squared` that can be used to calculate the reduced chi-squared value

<!-- grade_id: problem_3_function, points: 0, grade: false, solution: true -->
```python
def chi_squared(observed, expected, std):
    """Returns the reduced chi-squared value for input array data."""
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_function_tests, points: 1.5, grade: true, solution: false -->
```python
# These visible tests with fake data should work
from nose.tools import ok_, assert_equal

# Fake data
obs1 = np.array([1.1, 2.9, 2.6, 3.5, 5.7, 2.8])
exp1 = np.array([1.5, 2.4, 3.6, 1.5, 6.7, 2.6])
std1 = np.array([0.5, 1.4, 0.6, 1.5, 0.7, 1.0])

obs2 = np.array([1.8, 2.3, 2.2, 3.9, 5.5, 2.4])
exp2 = np.array([1.2, 2.0, 3.9, 1.7, 6.1, 2.9])
std2 = np.array([0.2, 0.4, 1.6, 2.5, 0.5, 0.3])

# Fake goodness-of-fit values
cs1 = chi_squared(obs1, exp1, std1)
cs2 = chi_squared(obs2, exp2, std2)

# Print results
print("Goodness-of-fit for dataset 1: {0:.4f}.".format(cs1))
print("Goodness-of-fit for dataset 2: {0:.4f}.".format(cs2))

# Check that the chi-squared function works
assert_equal(round(cs1, 4), 1.2340)
assert_equal(round(cs2, 4), 2.6139)
```

## Part 2: Calculating goodness-of-fit (*2 points*)

Next, we can calculate the goodness-of-fit for each point along our line and the corresponding observed seasonal mean temperaure. We will do this by first creating arrays of expected temperature values for each season and time period, and then calculating the goodness-of-fit between those expected values and the observed temperatures.

- Create 4 arrays that contain the expected temperature value for each year based on the equation for the regression lines calculated in Problem 1
    - The variables `exp_winter` and `exp_summer` can be used for all years, while `exp_winter20` and `exp_summer20` can be used for the last 20 years
- Calculate the goodness-of-fit between the expected values and the observed mean seasonal temperatures
    - The variables `gof_winter` and `gof_summer` can be used for all years, while `gof_winter20` and `gof_summer20` can be used for the last 20 years

<!-- grade_id: problem_3_values, points: 0, grade: false, solution: true -->
```python
# Expected temperatures based on least-squares regression in Problem 1
exp_winter = None
exp_summer = None
exp_winter20 = None
exp_summer20 = None

# Goodness-of-fit values
gof_winter = None
gof_summer = None
gof_winter20 = None
gof_summer20 = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_3_values_tests1, points: 1, grade: true, solution: false -->
```python
# This test should work
print('My last expected winter temperature is {0:.4f} and should be -3.4019.'.format(exp_winter[-1]))

```

<!-- grade_id: problem_3_values_tests2, points: 1, grade: true, solution: false -->
```python
# This test should work
print('My summer (all years) goodness-of-fit is {0:.4f} and should be 0.0653.'.format(gof_summer))

```

```python
# Run this cell to print all of your calculated goodness-of-fit values
print('My winter goodness-of-fit values:\n    All years:     {0:.4f}\n    Last 20 years: {1:.4f}\n'.format(gof_winter, gof_winter20))
print('My summer goodness-of-fit values:\n    All years:     {0:.4f}\n    Last 20 years: {1:.4f}'.format(gof_summer, gof_summer20))
```

# Summary questions for Exercise 2 (*2 points; 0.5 points each*)

Answer the questions below in a few sentences.

1. Based on the least-squares regressions calculated in Problem 1, how have summer and winter temperatures changed in Helsinki over the time 1953-2019? Is the trend similar for both seasons? What about when you consider only the data over the last 20 years? Do you have any first-hand experience that relates to these temperature trends? If so, what?
2. Based on the calculation of the correlation coefficients in Problem 2, how well does it seem the data are fit by a line? Are there any significant differences between seasons or time periods in the calculated $r$ values?
3. Looking at the goodness-of-fit values calculated in Problem 3, how well does it appear the data are fit by a regression line? Do you see differences between seasons ot time period considered? 
4. Comparing the calculated $\chi^{2}$ values from Problem 3 to the $r$ values in Problem 2, what differences do you see in the numbers between seasons and/or time periods? Can you think of reasons why the $r$ and $\chi^{2}$ values might differ in terms of their relative values in each season/time period? How are the $r$ and $\chi^{2}$ values different in terms of what they measure?

<!-- grade_id: summary_questions, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE